# Coursework 3 - CAL (Assessed) #

This coursework corresponds to Calculus (CAL).

### Marks 
- There are 3 subquestions, each worth 30 marks. 
- The marks are divided amongst different tests in each sub question.
- The number of marks asssociated with an automated test-cell is indicated in a comment at the top of that cell. 

#### How to Answer
Write your code in the "answer" cell where it says "# YOUR CODE HERE"
and remove both the line "# YOUR CODE HERE" and the line "raise NotImplementedError()"

**Important**: Do not change the name or remove the existing function, and write all your code "inside" the existing function, i.e. with at least one-tab's indentation.

When you have written your code you should evaluate the cell and then evaluate the "Test" cells following it. If you don't see any errors then your code is (hopefully) correct and you can go on to the next question. 

If you see some errors then you have to go back and edit your code and try to fix the "bug". 

For all questions you are allowed to use any imported modules that you like.

All plots should be done 'inline', in other words, do not use "popup" windows. With the default settings you just need to call 'plt.show()' at the end. 


## Question 2 [CAL]

The sequence of functions $F_n(x,y)$ for ($n\ge0$) is defined by the recurrence formula

$$
F_0(x,y)=1,\quad F_1(x,y)=\frac{x}{2} - y^2
$$
and 
$$
2n^2 F_{n+1}(x,y) = 3nxyF_n(x,y) - (2n+1)F_{n-1}(x,y),\quad n\ge 1.
$$

**(a) [5 marks]** Write a recursive function (i.e. a function that calls itself) `evalFn` which evaluates the function $F_n$ with the arguments $x$ and $y$.
- The function should return a float if the input is a scalar;
- If the input x and y are both two-dimensional NumPy arrays then the output should also be a two-dimensional NumPy array;
- If the input are numpy arrays of different shapes it should raise a ValueError.

Hint: You might find it useful to replace $n+1$ by $m$ and solve for $F_m$.

<!-- grade_id: cell-1f0a28112da0f41a, points: 0, grade: false, solution: true -->
```python
import numpy as np
def evalFn(n,x,y):
    # YOUR CODE HERE
    xyarray = False
    if type(x) == np.ndarray and type(y) == np.ndarray:
        xyarray = True
        if x.shape != y.shape:
            raise ValueError()
    
    if n == 0:
        if xyarray == True:
            return np.full(x.shape,1)
        else:
            return 1
    if n == 1:
        return ((0.5*x)-(y**2))
    else:
        return (3*(n-1)*x*y*(evalFn((n-1),x,y))-((2*(n-1))+1)*(evalFn((n-2),x,y)))/(2*((n-1)**2))
    raise NotImplementedError()
```

<!-- grade_id: cell-e904974906967df0, points: 3, grade: true, solution: false -->
```python
# 3 marks
# Test that 'evalFn' returns a float if a scalar is input.
from nose.tools import assert_equal
assert_equal(type(evalFn(1,1,1)),float)
```

<!-- grade_id: cell-8611d0e9fad64f34, points: 3, grade: true, solution: false -->
```python
# 3 mark
# Test that 'evalFn' outputs the correct values for some values of n
from nose.tools import assert_equal,assert_almost_equal
assert_almost_equal(evalFn(4,1,1),0.1875,delta=1e-10)
```

<!-- grade_id: cell-4f586f3852a5a124, points: 3, grade: true, solution: false -->
```python
# 3 mark
# Hidden tests that 'evalFn' outputs the correct values for some other values of n
from nose.tools import assert_equal,assert_almost_equal
assert_equal(evalFn(1,3,100),-9998.5)
```

<!-- grade_id: cell-17961f3c132d22ae, points: 3, grade: true, solution: false -->
```python
# 3 mark
# Hidden tests that 'evalFn' outputs the correct values for some other values of n
from nose.tools import assert_equal,assert_almost_equal
```

<!-- grade_id: cell-6865d5b69b8cd92a, points: 3, grade: true, solution: false -->
```python
# 3 marks
# Test that the 'evalFn' returns a numpy array of correct shape if the inputs are numpy arrays.
from nose.tools import assert_equal,assert_almost_equal,assert_raises,assert_true
import numpy as np
x = np.linspace(0,1,2)
y = np.linspace(0,1,3)
X,Y = np.meshgrid(x,y)
assert_equal(type(evalFn(0,X,Y)),np.ndarray)

```

<!-- grade_id: cell-d1323e2fdffa7d44, points: 3, grade: true, solution: false -->
```python
# 3 marks
from nose.tools import assert_equal,assert_almost_equal,assert_raises,assert_true
import numpy as np
```

<!-- grade_id: cell-942128e3f1d2406e, points: 3, grade: true, solution: false -->
```python
# 3 marks
# Test that the 'evalFn' returns a numpy array of correct shape if the inputs are numpy arrays.
from nose.tools import assert_equal,assert_almost_equal,assert_raises,assert_true
import numpy as np
x = np.linspace(0,1,2)
y = np.linspace(0,1,3)
X,Y = np.meshgrid(x,y)
assert_equal(evalFn(1,X,Y).shape,(3,2),"Function should return a 3 x 2 grid!")
```

<!-- grade_id: cell-e47879c81586be00, points: 3, grade: true, solution: false -->
```python
# 3 mark
# Test that 'evalFn' returns the correct values for some input vectors.
import numpy as np
x = np.linspace(0,1,2)
y = np.linspace(0,1,3)
X,Y = np.meshgrid(x,y)
assert_true( (evalFn(0,X,Y)==np.ones((3,2))).all())
```

<!-- grade_id: cell-6e1922c2e2829a04, points: 3, grade: true, solution: false -->
```python
# 3 mark
# Test that 'evalFn' returns the correct values for some input vectors.
import numpy as np
x = np.linspace(0,1,2)
y = np.linspace(0,1,3)
X,Y = np.meshgrid(x,y)
assert_almost_equal(evalFn(3,X,Y)[2,1],-1.375,delta=1e-6)
```

<!-- grade_id: cell-7ef7b518684fbde8, points: 3, grade: true, solution: false -->
```python
# 3 marks
# Test that 'evalFn' returns the correct values for some input vectors.
import numpy as np
```

**(b)**
In addition to working numerically we can also use SymPy to study the functions $F_n$ symbolically and obtain explicit expressions.

Use `SymPy` to write a function `symbolicFn` which gives an expression for $F_n(x,y)$ as a polynomial in $x$ and $y$ of the form 

$$
\sum_{i,j} a(i,j)x^i y^j,
$$
where $a(i,j)$ are simply some coefficients. An example of such an expression is 
$$
 xy + 2x^4y^2 + y^3
$$
where the coefficients would be $a(1,1)=1,a(4,2)=2,a(0,3)=1$ and all other $a(i,j)$ are zero. 

The function should take a non-negative integer $n$ and symbolic values for $x$ and $y$ (as given by e.g. `x, y = sympy.symbols('x y')`) as input and output a polynomial in $x$ and $y$ of the type `sympy.Poly` with domain the rational numbers (i.e. `sympy.QQ`).
- If n is not a non-negative integer the function should raise a ValueError
- If x or y are not symbolic expressions (i.e. of type `sympy.core.symbol.Symbol`) then a ValueError should be raised.

Hints: you might find the functions .as_poly() and sympy.Poly useful.

<!-- grade_id: cell-7c81740c09f6b81b, points: 0, grade: false, solution: true -->
```python
import sympy
def symbolicFn(n,x,y):
    # YOUR CODE HERE
    if n < 0:
        raise ValueError()
    if type(x) != sympy.core.symbol.Symbol and type(y) != sympy.core.symbol.Symbol:
        raise ValueError()
    
    if n == 0:
        return sympy.Poly(1,x,y,domain=sympy.QQ)
    elif n == 1:
        return sympy.Poly((0.5*x)-(y**2),x,y,domain=sympy.QQ)
    else:
        return sympy.poly((3*(n-1)*x*y*symbolicFn(n-1,x,y) - (2*(n-1) + 1)*symbolicFn(n-2,x,y))/(2*((n-1)**2)))
    
    raise NotImplementedError()
```

```python
print(symbolicFn(3,x,y))
```

Output:
```
Poly(9/16*x**3*y**2 - 9/8*x**2*y**4 - 9/8*x*y - 5/16*x + 5/8*y**2, x, y, domain='QQ')
```

<!-- grade_id: cell-35a33229368f8986, points: 6, grade: true, solution: false -->
```python
# 6 marks
# Test that the above algorithm provides reasonable outcome.
from nose.tools import assert_equal
import sympy
x, y = sympy.symbols('x y')
# Check that we get back the constant 1 as a polynomial in x and y (even though the only power is x^0y^0)
assert_equal(symbolicFn(0,x,y),sympy.Poly(1,x,y,domain=sympy.QQ)) 

```

<!-- grade_id: cell-d4d1b00bedeca95c, points: 6, grade: true, solution: false -->
```python
# 6 marks
# Test that we handle wrong input
from nose.tools import assert_raises
import sympy
x, y = sympy.symbols('x y')
assert_raises(ValueError,symbolicFn,0,1,1)
```

<!-- grade_id: cell-981ab513ffc7b528, points: 6, grade: true, solution: false -->
```python
# 6 marks
# Check that we get back the correct polynomial F_1
from nose.tools import assert_equal
import sympy
x, y = sympy.symbols('x y')
# assert_equal(symbolicFn(1,x,y),sympy.Poly(x/2-y**2,x,y))
assert_equal(symbolicFn(1,x,y),sympy.Poly(-y**2+x/2,x,y,domain=sympy.QQ))

```

<!-- grade_id: cell-fde1d11ca4a539a7, points: 6, grade: true, solution: false -->
```python
# 6 marks
# Check that we get back the correct polynomial F_3
from nose.tools import assert_equal
import sympy
x, y = sympy.symbols('x y')
assert_equal(symbolicFn(3,x,y),sympy.Poly(9/16*x**3*y**2 - 9/8*x**2*y**4 - 9/8*x*y - 5/16*x + 5/8*y**2,x,y,domain=sympy.QQ))
```

<!-- grade_id: cell-f9e38d1a75e8b94d, points: 6, grade: true, solution: false -->
```python
# 6 marks
from nose.tools import assert_equal
import sympy
```

**(c)**
Consider the following function of two variables:
$$
G(x,y) = 1 + x^2 y^2 
$$
Use SymPy for the following:
1.  Write a function `S` with a parameter $x$ to evaluate the function $S(x)$ given by the following integral:
$$
S(x) = \int_{x^2}^{1} G(x,y) dy
$$
as a symbolic expression in $x$.
2. Write a function `D(x,n)` which takes as arguments a symbol `x` and an integer $n$ to evaluate the $n$-th derivative of $S(x)$ with respect to $x$.
3. Finally write a function `T` (without arguments) to evaluate the integeral $\int_{0}^{1} S(x)dx$.

The functions $S(x)$ and $D(x,n)$ should both return polynomials of the type `sympy.Poly`, variable `x` and domain `sympy.QQ`.

Hint: 
- Note that the the output of "sympy.integrate" is not a SymPy polynomial.
- Note that the function "sympy.diff" does not accept a SymPy polynomial as input. 
- To convert input and output you might find the functions `.as_expr` and `as_poly` useful.


<!-- grade_id: cell-0551ae10a9329aa5, points: 0, grade: false, solution: true -->
```python
import sympy
def S(x):
    # YOUR CODE HERE
    G = 1+(x**2)*(y**2)
    S = (sympy.integrate(G,y)).as_expr
    return sympy.Poly((S.subs({y:1}) - S.subs({y:x**2})), domain=sympy.QQ)
    raise NotImplementedError()

def D(x,n):
    # YOUR CODE HERE
    
    raise NotImplementedError()

def T(a):
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
x, y = sympy.symbols('x,y')
G = 1+(x**2)*(y**2)
S = sympy.integrate(G,y)
type(G)
type(S)
sympy.Poly((S.subs({y:1}) - S.subs({y:x**2})), domain=sympy.QQ)
```

Output:
```
Poly(-1/3*x**8 - 2/3*x**2 + 1, x, domain='QQ')
```

<!-- grade_id: cell-8aeacee668b6e1d8, points: 5, grade: true, solution: false -->
```python
# Marks: 6
from nose.tools import assert_equal
import sympy
x, y = sympy.symbols('x,y')
assert_equal(S(x),sympy.Poly(-x**8/3 - 2*x**2/3 + 1,domain=sympy.QQ))

```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-216-4c37457ca262> in <module>
      3 import sympy
      4 x, y = sympy.symbols('x,y')
----> 5 assert_equal(S(x),sympy.Poly(-x**8/3 - 2*x**2/3 + 1,domain=sympy.QQ))

TypeError: 'Add' object is not callable
```

<!-- grade_id: cell-f7d7cf3883847bda, points: 5, grade: true, solution: false -->
```python
# Marks: 6
from nose.tools import assert_equal
x, y = sympy.symbols('x,y')
assert_equal(D(x,1),sympy.Poly(-8/3*x**7 - 4/3*x, x, domain=sympy.QQ))
```

<!-- grade_id: cell-6b2427e752a17691, points: 5, grade: true, solution: false -->
```python
# Marks: 6
from nose.tools import assert_equal
```

<!-- grade_id: cell-3d5d93b52065bca7, points: 5, grade: true, solution: false -->
```python
# Marks: 6
from nose.tools import assert_equal
x, y = sympy.symbols('x,y')
assert_equal(D(x,9),sympy.Poly(0,x,domain=sympy.QQ))
```

<!-- grade_id: cell-f49b5de60e53b495, points: 5, grade: true, solution: false -->
```python
# Marks: 6
```

<!-- grade_id: cell-c6a709dbc1b13f7c, points: 5, grade: true, solution: false -->
```python
# Marks: 6
from nose.tools import assert_equal
x, y = sympy.symbols('x,y')
assert_equal(T(sympy.Rational(1,2)),sympy.Rational(6527,13824))
```