# Coursework 3 - LMA (Assessed) #

This coursework corresponds to Linear Mathematics (LMA)

### Marks 
- There are 3 subquestions, each worth 30 marks. 
- The marks are divided amongst different tests in each sub question.
- The number of marks asssociated with an automated test-cell is indicated in a comment at the top of that cell. 

#### How to Answer
Write your code in the "answer" cell where it says "# YOUR CODE HERE"
and remove both the line "# YOUR CODE HERE" and the line "raise NotImplementedError()"

**Important**: Do not change the name or remove the existing function, and write all your code "inside" the existing function, i.e. with at least one-tab's indentation.

When you have written your code you should evaluate the cell and then evaluate the "Test" cells following it. If you don't see any errors then your code is (hopefully) correct and you can go on to the next question. 

If you see some errors then you have to go back and edit your code and try to fix the "bug". 

For all questions you are allowed to use any imported modules that you like.

All plots should be done 'inline', in other words, do not use "popup" windows. With the default settings you just need to call 'plt.show()' at the end. 


## Introduction

The Gram-Schmidt process is an algorithm which can be used to construct a set of vectors $\{\mathbf{w}_i\}$ which are **orthonormal**, in other words they are pairwise orthogonal to each other and have lengths equal to 1. This condition can be  expressed in terms of the standard Euclidean scalar (or "dot") product as 

$$
\mathbf{w}_i \cdot \mathbf{w}_j = 
\begin{cases}
1 & \text{if } i = j,\\
0 & \text{if } i\ne j.
\end{cases}
$$

If the input to the Gram-Schmidt process is a set of $n$ linearly independent vectors $\{ \mathbf{v}_i \mid i=1,\ldots,n\}$ in $\mathbb{R}^{n}$, then the output will be a set of $n$ orthonormal vectors $\{\mathbf{w}_i \mid i=1,\ldots,n\}$ in $\mathbb{R}^{n}$.

In general, if the input vectors are not all linearly independent, then the output will be a set of $n$ vectors containing an orthonormal basis for the linear span of the vectors $\{ \mathbf{v}_i\}$ together with a number of zero-vectors.

The method is summarised by the following instructions:
1. Construct an orthogonal set of vectors $\{\mathbf{w}_i\}$ using the formula: $\mathbf{w}_1=\mathbf{v}_1$ and 

$$
\mathbf{w}_i = \mathbf{v}_i - \sum_{j=1}^{i-1} \frac{\mathbf{v}_i\cdot \mathbf{w}_j}{\| \mathbf{w}_j\|^2}\mathbf{w}_j,\quad i>1.
$$

2. If $\mathbf{w}_i$ from the previous step is not the zero vector $\mathbf{0}$, normalise it to have length 1 by setting 

$$
\mathbf{w}_i = \frac{\mathbf{w}_i}{\|\mathbf{w}_i \|}.
$$


**(a) [45 marks]**
Write a function `gram_schmidt_np` which takes as input a list of length $n$ containing $n$ NumPy arrays, $\{\mathbf{v}_i\}$, each of length $n$ and then uses the Gram-Schmidt process to construct an orthonormal set of vectors $\{\mathbf{w}_i\}$, which is then returned (in order).

- Your function should check that the input V is list of arrays of the same length and raise a ValueError if not.
- In the normalisation step you can check if $w_i$ has norm less than $1\epsilon=1.19\cdot 10^{-7}$ instead of checking if it is exactly 0 and in this case just replace it by the zero vector instead of normalising it. 

Note: The number $\epsilon$ is sometimes called "machine-epsilon" and can be obtained in NumPy for single precision floating point numbers as `np.finfo('float32').eps`)



<!-- grade_id: cell-2c038394d0ebff43, points: 0, grade: false, solution: true -->
```python
import numpy as np
def gram_schmidt_np(V):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0eb1f34611a7be75, points: 10, grade: true, solution: false -->
```python
# Marks: 6
from nose.tools import assert_almost_equal
import numpy as np
# Check that it works for a specific matrix V of linearly independent vectors.
V1 = [np.array([1,3]), np.array([2,4])]
W1 = [np.array([ 0.31622777,  0.9486833 ]), np.array([ 0.9486833 , -0.31622777])]
assert_almost_equal(np.linalg.norm(np.array(W1)-gram_schmidt_np(V1)),0,delta=1e-8)
```

<!-- grade_id: cell-e7e4ba4c4a50dab2, points: 10, grade: true, solution: false -->
```python
# Marks: 3
from nose.tools import assert_almost_equal
import numpy as np
# Check that it works for a case when vectors are linearly dependent
V2 = [np.array([1,2]),np.array([2,4])]
W2 = [np.array([0.4472136 , 0.89442719]), np.array([0., 0.])]
assert_almost_equal(np.linalg.norm(np.array(W2)-gram_schmidt_np(V2)),0,delta=1e-8)
```

<!-- grade_id: cell-7304f69f36376e3e, points: 10, grade: true, solution: false -->
```python
 Marks: 3
from nose.tools import assert_almost_equal
import numpy as np
V3 = [np.array([1,4,7]),np.array([2,5,8]),np.array([3,6,10])]
W3 = [np.array([0.12309149, 0.49236596, 0.86164044]), np.array([ 0.90453403,  0.30151134, -0.30151134]), 
      np.array([ 0.40824829, -0.81649658,  0.40824829])]
             
assert_almost_equal(np.linalg.norm(np.array(W3)-gram_schmidt_np(V3)),0,delta=1e-8)

```

<!-- grade_id: cell-ade454bdd04d761c, points: 10, grade: true, solution: false -->
```python
# Marks: 3
from nose.tools import assert_almost_equal
import numpy as np
```

<!-- grade_id: cell-d23c23d8b7f903cd, points: 3, grade: true, solution: false -->
```python
# Marks: 3
# Check that errors are raised for wrong input (vectors not of same length).
from nose.tools import assert_raises
assert_raises(ValueError,gram_schmidt_np,[np.array([1,2]),np.array([1,5,1]),np.array([7,8])])
```

<!-- grade_id: cell-96ee3f8ec2cfbc48, points: 2, grade: true, solution: false -->
```python
# Marks: 3
# Check that errors are raised for wrong input (not a list)
from nose.tools import assert_raises
assert_raises(ValueError,gram_schmidt_np,np.array([1,2]))
```

Instead of using NumPy for working with vectors it is also posible to use SymPy.

**(b) [45 marks]** 

Write a function `gram_schmidt_sp` which takes as argument an $n \times n$ square matrix $V$. The columns of $$V$$ represent the set of $n$-dimensional vectors $\{\mathbf{v}_i\}$ and the function should use the Gram-Schmidt process to construct an orthonormal set of vectors $\{\mathbf{w}_i\}$, which are then returned in the form of rows of a matrix W.

- The input matrix V can be assumed to be of type sympy.Matrix 
- Your function should check that the input V is a square matrix of type sympy.Matrix, and raise a ValueError if not.
- The output matrix W should also be of type sympy.Matrix.


Note: while the algorithm is the same as above you will need to figure out for yourself how to extract columns from a SymPy matrix (the columns will themselves be $n \times 1$ matrices), how to take the dot product and norm of a SymPy matrix and how to join rows to make a new SymPy matrix. 
(Hint: It is often easier to join columns and then take the transpose. The function sympy.Matrix.hstack can be useful.)



<!-- grade_id: cell-a26a550fe2ca1c0c, points: 0, grade: false, solution: true -->
```python
import sympy
def gram_schmidt_sp(V):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a61ea02026825428, points: 10, grade: true, solution: false -->
```python
# Marks: 10
from nose.tools import assert_equal
import sympy
# Test that it works for a specific matrix V
V1 = sympy.Matrix([[1,2],[3,4]])
W1 = sympy.Matrix([[sympy.sqrt(10)/10, 3*sympy.sqrt(10)/10], [3*sympy.sqrt(10)/10, -sympy.sqrt(10)/10]])
assert_equal(W1-gram_schmidt_sp(V1),sympy.zeros(2,2))
```

<!-- grade_id: cell-75aa51ccb385b6d8, points: 10, grade: true, solution: false -->
```python
# Marks: 10
from nose.tools import assert_equal
import sympy
# Check that it works for a case when vectors that are linearly dependent
V2 = sympy.Matrix([[1,2],[2,4]])
W2 = sympy.Matrix([[sympy.sqrt(5)/5, 2*sympy.sqrt(5)/5], [0, 0]])
assert_equal(W2-gram_schmidt_sp(V2),sympy.zeros(2,2))

```

<!-- grade_id: cell-6b541eaac550c772, points: 10, grade: true, solution: false -->
```python
# Marks: 10
from nose.tools import assert_equal
import sympy
from sympy import sqrt
# Hidden test that it works for another matrix input.
V3 = sympy.Matrix([[1,2,3],[4,5,6],[7,8,10]])
W3=sympy.Matrix([[sqrt(66)/66, 2*sqrt(66)/33, 7*sqrt(66)/66], 
                 [3*sqrt(11)/11, sqrt(11)/11, -sqrt(11)/11], 
                 [sqrt(6)/6, -sqrt(6)/3, sqrt(6)/6]])
assert_equal(W3-gram_schmidt_sp(V3),sympy.zeros(3,3))
```

<!-- grade_id: cell-362bc19198870bc0, points: 10, grade: true, solution: false -->
```python
# Marks: 10
from nose.tools import assert_equal
import sympy
from sympy import sqrt
# Hidden test that it works for another matrix input.
```

<!-- grade_id: cell-51db1d9629d2a80f, points: 2, grade: true, solution: false -->
```python
# Marks: 2
from nose.tools import assert_raises
import sympy
# Check that errors are raised for wrong input.
import sympy
assert_raises(ValueError,gram_schmidt_sp,1)
```

<!-- grade_id: cell-289e3fed3ed188bb, points: 3, grade: true, solution: false -->
```python
# Marks: 3
from nose.tools import assert_raises
import sympy
# Check that errors are raised for wrong input.
import sympy
assert_raises(ValueError,gram_schmidt_sp,sympy.Matrix([[1,2],[3,4],[5,6]]))
```