# Coursework 1 - LMA  #

This question is associated with the core module Linear Mathematics (LMA).

### Marks 
- There are 3 subquestions, each is worth 30 marks.
- The number of marks asssociated with test-cell is indicated in a comment at the top of that cell.
- The number of marks associated to a manually marked taks. (like drawing a graph) is indicated in the text immediately above this cell. 

#### How to Answer
Write your code in the "answer" cell where it says "# YOUR CODE HERE"
and remove both the line "# YOUR CODE HERE" and the line "raise NotImplementedError()"

**Important**: Do not change the name or remove the existing function and write all your code "inside" the existing function, i.e. with at least one-tab's indentation.


When you have written your code you should evaluate the cell and then evaluate the "Test" cells following it. If you don't see any errors then your code is (hopefully) correct and you can go on to the next question. 

If you see some errors then you have to go back and edit your code and try to fix the "bug". 

## Introduction

For this question we will deal with geometry of the plane and we will use tuples to denote points and vectors. 
More precisely, if $\bf{v}$ is a vector or point in the plane we represent it by the tuple of coordinates $\bf{v}=(x,y)$. A similar representation can of course be used for vectors in arbitrary dimensions. 

We know that a straight line $L$ in the plane is uniquely determined by either of the following:

- Two different points $P$ and $Q$ that lie on the line. 
- A point $P$ on the line together with a direction vector $\bf{v}$

and we will make use of both of these definitions. 

**a)**
Write a function that computes the scalar product of two vectors of arbitrary but equal length.  
The function should Raise a ValueError error if the input are not tuples of equal length containing integers or floats. 

<!-- grade_id: cell-928b75d38d6fe5f1, points: 0, grade: false, solution: true -->
```python
def scalar_product(v1,v2):
    r"""
    Calculate the scalar (dot) product between two vectors, v1 and v2, of arbitrary length.
    """
    # YOUR CODE HERE
    assert type(v1) == tuple and type(v2) == tuple and len(v1) == len(v2)
    v_dot = 0
    for x in range(len(v1)):
        v_dot += v1[x]*v2[x]
    
    return v_dot
    raise NotImplementedError()

```

```python

```

<!-- grade_id: cell-ce688c52a91885f5, points: 5, grade: true, solution: false -->
```python
# [5 marks]
from nose.tools import assert_equal
assert_equal(scalar_product((1,2,0),(-2,1,-1)),0)
```

<!-- grade_id: cell-0799593c3983ba38, points: 5, grade: true, solution: false -->
```python
# [5 marks]
from nose.tools import assert_equal
assert_equal(scalar_product((1,2,1),(-2,1,-1)),-1)
```

<!-- grade_id: cell-f475927cdf9d0677, points: 5, grade: true, solution: false -->
```python
# [5 marks]
from nose.tools import assert_equal
# This cell contains hidden tests.

```

<!-- grade_id: cell-fc7bf5c03189a60d, points: 5, grade: true, solution: false -->
```python
# [5 marks]
from nose.tools import assert_equal
# This cell contains hidden tests.

```

<!-- grade_id: cell-294594a5f5da639e, points: 5, grade: true, solution: false -->
```python
# [5 marks]
from nose.tools import assert_raises
# Check that a ValueError is raised if the input is not tuples
assert_raises(ValueError,scalar_product,1,2) # not tuples as inout
```

Output:
```
~\anaconda3\lib\unittest\case.py in handle(self, name, args, kwargs)
    200                 self.obj_name = str(callable_obj)
    201             with self:
--> 202                 callable_obj(*args, **kwargs)
    203         finally:
    204             # bpo-23890: manually break a reference cycle

<ipython-input-1-7b358f537f12> in scalar_product(***failed resolving arguments***)
      4     """
      5     # YOUR CODE HERE
----> 6     assert type(v1) == tuple and type(v2) == tuple and len(v1) == len(v2)
      7     v_dot = 0
      8     for x in range(len(v1)):

AssertionError:
```

<!-- grade_id: cell-d70098b1237e1380, points: 5, grade: true, solution: false -->
```python
# [5 marks]
from nose.tools import assert_raises
# This cell contains hidden tests.
```

**b)** Write a program that takes as input two points $P$ and $Q$ in the plane and returns a vector that is perpendicular to the line through $P$ and $Q$ 


<!-- grade_id: cell-4d0f9c5c5f7a0541, points: 0, grade: false, solution: true -->
```python
import math
def vector_perpendicular(P,Q):
    """
    Find a vector that is perpendicular to the line between two points P and Q in the plane.  
    """
    # YOUR CODE HERE
    PQ = (Q[0]-P[0],Q[1]-P[1])
    for a in range(5000):
        for b in range(5000):
            if a != 0 or b != 0:
                if PQ[0]*a + PQ[1]*b == 0:
                    return (a,b)
                if PQ[0]*(-a) + PQ[1]*b == 0:
                    return ((-a),b)
                if PQ[0]*a + PQ[1]*(-b) == 0:
                    return (a,(-b))
                if PQ[0]*(-a) + PQ[1]*(-b) == 0:
                    return ((-a),(-b))
    raise NotImplementedError()
```

```python
# This function is used in the tests below to check whether or not two vectors are orthogonal.
# If you run the tests on your own make sure to run this cell first. 
def are_orthogonal(v,w):
    """
    Simple check to see if the two vectors v and w are orthogonal up to precision 1e-10.
    """
    return abs(v[0]*w[0]+v[1]*w[1])<1e-10
```

<!-- grade_id: cell-e72bd17faa23394d, points: 5, grade: true, solution: false -->
```python
# [5 marks]
# Test that checks the function produces a vector that is orthogonal to the line.
from nose.tools import assert_equal,assert_almost_equal
assert_equal( are_orthogonal(vector_perpendicular((0,1),(1,1)),(1,0)),True) 

```

<!-- grade_id: cell-52b3bb6c6f56aee9, points: 5, grade: true, solution: false -->
```python
# [5 marks]
# Test that checks the function produces a vector that is orthogonal to the line.
from nose.tools import assert_equal,assert_almost_equal
assert_equal( are_orthogonal(vector_perpendicular((1,2),(1,3)),(0,1)),True)
```

<!-- grade_id: cell-4acea58acdf1b264, points: 5, grade: true, solution: false -->
```python
# [5 marks]
# Test that checks the function produces a vector that is orthogonal to the line.
from nose.tools import assert_equal,assert_almost_equal
assert_equal( are_orthogonal(vector_perpendicular((1,0),(1,1)),(0,1)),True)

```

<!-- grade_id: cell-6d8b35c7a35fa33b, points: 5, grade: true, solution: false -->
```python
# [5 marks]
# Test that checks the function produces a vector that is orthogonal to the line.
from nose.tools import assert_equal,assert_almost_equal
# This cell contains hidden tests.
```

<!-- grade_id: cell-321bd15b150dd805, points: 5, grade: true, solution: false -->
```python
# [5 marks]
from nose.tools import assert_equal,assert_almost_equal
# This cell contains hidden tests.
```

<!-- grade_id: cell-5ce71af93ba476c0, points: 5, grade: true, solution: false -->
```python
# [5 marks]
from nose.tools import assert_equal,assert_almost_equal
# This cell contains hidden tests.
```

**c)** 
Write a function called `intersection_lines` that computes the point of intersection between two lines $L_1$ and $L_2$ in the plane. Both lines are given in the format $(P,v)$ where $P$ is a point and $v$ a vector. We use tuples of length 2 to represent both the point and the vector so the input parameter is a tuple of length 2 with each element being a tuple of length 2.  

The function should return the point of intersection as a tuple (x,y) where x and y are floats. 

If the lines are parallell  so that they don't intersect then the function should raise a ValueError.

**Hint for the calculation**
If we are given the lines in parameter format

$$
L_1: s\mapsto  P_1 + s\mathbf{v_1} \quad L_2: t\mapsto P_2 + t\mathbf{v_2}
$$
and want to find their intersection point then we need to solve the linear system of equations
 $L_1(s)=L_2(t)$ for $s$ and $t$.
If $P_i=(a_i,c_i)$ and $\mathbf{v}_i = (b_i,d_i)$ then the system of equations take the form:
$$
   a_1 + b_1 s = a_2 + b_2 t\\
   c_1 + d_1 s = c_2 + d_2 t
$$
and the aim is to find $s$ and $t$ in terms of the other coefficients. 
You will also need to find out how you can tell from the coefficients in the linear equation if the two lines are parallel or not.

It will be very helpful to work out the equations on a piece of paper instead of trying to program it directly. 

<!-- grade_id: cell-447639baabe020c7, points: 0, grade: false, solution: true -->
```python
def intersection_lines(L1,L2):
    """
    Find the intersection point of the two lines L1 and L2 given by two tuples L1=(P1,v1) and L2=(P2,v2) 
    where the points P1, P2 and the vectors v1 and v2 are all represented by tuples of length 2. 
    
    If the vectors v1 and v2 are parallell then we raise an assertion error. 
    """
    # YOUR CODE HERE
    assert L1[1] != L2[1]
    if L1[1][0] == 0 and L2[1][0] == 0:
        assert 1==1
    if L1[1][1] == 0 and L2[1][1] == 0:
        assert 1==1
    if L1[1][0] != 0 and L1[1][1] != 0 and L2[1][0] != 0 and L2[1][1] != 0:
        assert L1[1][0]/L2[1][0] != L1[1][1]/L2[1][1]
    
    a1=L1[0][0]
    b1=L1[1][0]
    c1=L1[0][1]
    d1=L1[1][1]
    a2=L2[0][0]
    b2=L2[1][0]
    c2=L2[0][1]
    d2=L2[1][1]

        
    t=(a1*d1-a2*d1+b1*c2-b1*c1)/(b2*d1-b1*d2)
    s=(a2+b2*t-a1)/b1
    return ((float(a1+s*b1)),(float(c1+d1*s)))
    raise NotImplementedError()

```

<!-- grade_id: cell-8bfb239148e0fa59, points: 5, grade: true, solution: false -->
```python
# [5 marks]
from nose.tools import assert_equal
assert_equal(intersection_lines(((1,1),(1,1)),((1,1),(0,1))),(1.0,1.0))
   
```

<!-- grade_id: cell-b39645ad4022234e, points: 5, grade: true, solution: false -->
```python
# [5 marks]
from nose.tools import assert_equal
assert_equal(intersection_lines(((1,-2),(1,1)),((1,1),(0,1))),(1.0,-2.0))

```

<!-- grade_id: cell-d17b1866346a67bf, points: 5, grade: true, solution: false -->
```python
# [5 marks]
from nose.tools import assert_equal
# This cell contains hidden tests
```

<!-- grade_id: cell-172e2dad1bf9b89e, points: 5, grade: true, solution: false -->
```python
# [5 marks]
from nose.tools import assert_equal
assert_equal(intersection_lines(((-1,1),(-1,1)),((1,1),(1,1))),(0,0))
```

<!-- grade_id: cell-2e0f6de4ff56da0c, points: 5, grade: true, solution: false -->
```python
# [5 marks]
from nose.tools import assert_raises
assert_raises(ValueError,intersection_lines,((1,0),(0,1)),((1,1),(0,1)))
```

Output:
```
~\anaconda3\lib\unittest\case.py in handle(self, name, args, kwargs)
    200                 self.obj_name = str(callable_obj)
    201             with self:
--> 202                 callable_obj(*args, **kwargs)
    203         finally:
    204             # bpo-23890: manually break a reference cycle

<ipython-input-16-c9b29fb9cf47> in intersection_lines(***failed resolving arguments***)
      7     """
      8     # YOUR CODE HERE
----> 9     assert L1[1] != L2[1]
     10     if L1[1][0] == 0 and L2[1][0] == 0:
     11         assert 1==1

AssertionError:
```

<!-- grade_id: cell-11e17db173f166cb, points: 5, grade: true, solution: false -->
```python
# [5 marks]
from nose.tools import assert_raises
# This cell contains hidden tests

```