# Coursework 2 - LMA  #

This question is associated with the core module Analytical and Computational foundations (ACF).

### Marks 
- There are 3 subquestions, each is worth 30 marks.
- The number of marks asssociated with test-cell is indicated in a comment at the top of that cell.
- The number of marks associated to a manually marked taks. (like drawing a graph) is indicated in the text immediately above this cell. 

#### How to Answer
Write your code in the "answer" cell where it says "# YOUR CODE HERE"
and remove both the line "# YOUR CODE HERE" and the line "raise NotImplementedError()"

**Important**: Do not change the name or remove the existing function and write all your code "inside" the existing function, i.e. with at least one-tab's indentation.


When you have written your code you should evaluate the cell and then evaluate the "Test" cells following it. If you don't see any errors then your code is (hopefully) correct and you can go on to the next question. 

If you see some errors then you have to go back and edit your code and try to fix the "bug". 

## Introduction 

We have seen that it is possible to do linear algebra in NumPy with matrices described by 2-dimensional arrays. 
In this question we will write a function to perform Gaussian Elimination in Python using NumPy.

**(a) [30 Marks]**

Write a function that takes as input a Matrix A (given as a 2-dimensional NumPy-array) and two integers $i$ and $j$  and returns the entry in column $j$ with maximum absolute value amongst rows $i,i+1,..$ together with its row index (counted in the entire matrix).

In other words, if $A$ is an $r\times k$ matrix,  $A=(a_{mn})$ where $0\le m\le r-1$ and $0\le n\le k-1$, then the function should return 

$$
\max\{|a_{mj}|\; : {i \le m < r } \},
$$

together with the value of $m$ for which this maximum occurs. 
Note that indexing for rows and columns are counted from $0$ as usual. 

The function should raise a `ValueError` if the input $i$ is not the index of a column of $A$ or if $j$ 
is not the index of a row of $A$. 


<!-- grade_id: cell-8314a69148b9760e, points: 0, grade: false, solution: true -->
```python
def find_max_column_value(A,i,j):
    # YOUR CODE HERE
    
    if i > A.shape[0] or j > A.shape[1] or i < 0 or j < 0:
        raise IndexError()
    
    max_value = abs(A[i][j])
    m = i
    for x in range((i+1),A.shape[0]):
        if abs(A[x][j]) > abs(max_value):
            m = x
            max_value = abs(A[x][j])
    return (max_value,m)
    raise NotImplementedError()
```

<!-- grade_id: cell-7c913cef76d2617e, points: 5, grade: true, solution: false -->
```python
# Tests that the algorithm finds the correct value 
# Marks: 5
from nose.tools import assert_equal
import numpy as np
A = np.array([[ 9,  6,  7,  8,  1,  7,  2], [ 8,  2,  6,  5,  1,  5,  3], [ 7,  3,  1,  4,  5, 10,  1],
              [10,  5,  7,  5,  4,  6,  2], [ 5,  5,  2,  6,  4,  2,  7]])
assert_equal(find_max_column_value(A,0,0),(10,3))

```

<!-- grade_id: cell-5cafd6d3e799ea07, points: 5, grade: true, solution: false -->
```python
# Tests that the algorithm finds the correct value 
# Marks: 5
from nose.tools import assert_equal
import numpy as np
A = np.array([[ 9,  6,  7,  8,  1,  7,  2], [ 8,  2,  6,  5,  1,  5,  3], [ 7,  3,  1,  4,  5, 10,  1],
              [10,  5,  7,  5,  4,  6,  2], [ 5,  5,  2,  6,  4,  2,  7]])
assert_equal(find_max_column_value(A,2,4),(5,2))

```

<!-- grade_id: cell-ee59b89568fbcf45, points: 5, grade: true, solution: false -->
```python
# Tests that the algorithm raises IndexError 
# Marks: 5
from nose.tools import assert_raises
import numpy as np
A = np.array([[ 9,  6,  7,  8,  1,  7,  2], [ 8,  2,  6,  5,  1,  5,  3], [ 7,  3,  1,  4,  5, 10,  1],
              [10,  5,  7,  5,  4,  6,  2], [ 5,  5,  2,  6,  4,  2,  7]])
assert_raises(IndexError,find_max_column_value,A,-1,0)


```

<!-- grade_id: cell-ded2694e62fd425b, points: 5, grade: true, solution: false -->
```python
# Tests that the algorithm raises IndexError 
# Marks: 5
from nose.tools import assert_raises
import numpy as np
A = np.array([[ 9,  6,  7,  8,  1,  7,  2], [ 8,  2,  6,  5,  1,  5,  3], [ 7,  3,  1,  4,  5, 10,  1],
              [10,  5,  7,  5,  4,  6,  2], [ 5,  5,  2,  6,  4,  2,  7]])
assert_raises(IndexError,find_max_column_value,A,10,0)

```

<!-- grade_id: cell-5315e8ba04e5004a, points: 5, grade: true, solution: false -->
```python
# Tests that algorithm finds the correct value for another matrix
# Marks: 5
from nose.tools import assert_equal
import numpy as np
```

<!-- grade_id: cell-966636228d3e32ff, points: 5, grade: true, solution: false -->
```python
# Tests that the algorithm finds the correct value for another matrix
# Marks: 5
from nose.tools import assert_equal
import numpy as np
```

**(b) [30 Marks]**

Write two functions `swap_rows` and `swap_columns` that both take as input a matrix $A$ and two integers $i$ and $j$, and returns a copy of the matrix A where the  rows $i$ and $j$, or columns $i$ and $j$, are swapped respectively. 

The function should raise an `IndexError` if the input $i$ or $j$ is not the index of a row or column of $A$ respectively.

<!-- grade_id: cell-ae15f5b597bbd508, points: 0, grade: false, solution: true -->
```python
def swap_rows(A,i,j):
    # YOUR CODE HERE
    if i > A.shape[0]:
        raise IndexError()
    if i < 0:
        raise IndexError()
    if j > A.shape[0]:
        raise IndexError()
    if j < 0:
        raise IndexError()
        
    A1 = np.zeros((A.shape))
    for x in range(A.shape[0]):
        if x == min(i,j):
            A1[x] += (A[max(i,j)])
        elif x == max(i,j):
            A1[x] += (A[min(i,j)])
        else:
            A1[x] += A[x]
    
    return A1
    raise NotImplementedError()
```

<!-- grade_id: cell-22104d72b1eef0a1, points: 0, grade: false, solution: true -->
```python
def swap_cols(A,i,j):
    # YOUR CODE HERE
    AT = A.transpose()
    if i > AT.shape[0]:
        raise IndexError()
    if i < 0:
        raise IndexError()
    if j > AT.shape[0]:
        raise IndexError()
    if j < 0:
        raise IndexError()
        
    A1 = np.zeros((AT.shape))
    for x in range(AT.shape[0]):
        if x == min(i,j):
            A1[x] += (AT[max(i,j)])
        elif x == max(i,j):
            A1[x] += (AT[min(i,j)])
        else:
            A1[x] += AT[x]
    
    return A1.transpose()
    raise NotImplementedError()
```

<!-- grade_id: cell-e5c11fdb4d8c0942, points: 5, grade: true, solution: false -->
```python
# Check that the function 'swap_rows' work as expected
# Marks: 5
from nose.tools import assert_equal
import numpy as np
A = np.array([[ 9,  6,  7,  8,  1,  7,  2], 
              [ 8,  2,  6,  5,  1,  5,  3], 
              [ 7,  3,  1,  4,  5, 10,  1],
              [10,  5,  7,  5,  4,  6,  2], 
              [ 5,  5,  2,  6,  4,  2,  7]])
A1 = np.array([[ 7,  3,  1,  4,  5, 10,  1], 
               [ 8,  2,  6,  5,  1,  5,  3], 
               [ 9,  6,  7,  8,  1,  7,  2] ,
              [10,  5,  7,  5,  4,  6,  2], 
               [ 5,  5,  2,  6,  4,  2,  7]])

assert_equal( (swap_rows(A,0,2)==A1).all(),True)
```

<!-- grade_id: cell-8cf9aaff95b4f89c, points: 5, grade: true, solution: false -->
```python
# Check that the function swap_rows raises an IndexError when called with a row -1
from nose.tools import assert_raises
import numpy as np
A = np.array([[ 9,  6,  7,  8,  1,  7,  2], 
              [ 8,  2,  6,  5,  1,  5,  3], 
              [ 7,  3,  1,  4,  5, 10,  1],
              [10,  5,  7,  5,  4,  6,  2], 
              [ 5,  5,  2,  6,  4,  2,  7]])
assert_raises(IndexError,swap_rows,A,-1,1)
assert_raises(IndexError,swap_rows,A,5,-1)
```

<!-- grade_id: cell-c59173d62a3803c8, points: 5, grade: true, solution: false -->
```python
# Check that the function 'swap_rows' work as expected for a different matrix.
# Marks: 5
```

<!-- grade_id: cell-8610970d3fe60e4a, points: 5, grade: true, solution: false -->
```python
# Check that the function `swap_cols` work as expected
# Marks: 5
from nose.tools import assert_equal
import numpy as np
A = np.array([[ 9,  6,  7,  8,  1,  7,  2], 
              [ 8,  2,  6,  5,  1,  5,  3], 
              [ 7,  3,  1,  4,  5, 10,  1],
              [10,  5,  7,  5,  4,  6,  2], 
              [ 5,  5,  2,  6,  4,  2,  7]])
A2 = np.array([[ 8,  6,  7,  9,  1,  7,  2], 
               [ 5,  2,  6,  8,  1,  5,  3], 
               [ 4,  3,  1,  7,  5, 10,  1],
               [ 5,  5,  7,  10,  4,  6,  2], 
               [ 6,  5,  2,  5,  4,  2,  7]])

assert_equal( (swap_cols(A,0,3)==A2).all(),True)
```

<!-- grade_id: cell-79dbadf3393ef2e0, points: 5, grade: true, solution: false -->
```python
# Check that the function `swap_cols` raises IndexError when called with indexes out of bound.work as expected
# Marks: 5
from nose.tools import assert_raises
import numpy as np
A = np.array([[ 9,  6,  7,  8,  1,  7,  2], 
              [ 8,  2,  6,  5,  1,  5,  3], 
              [ 7,  3,  1,  4,  5, 10,  1],
              [10,  5,  7,  5,  4,  6,  2], 
              [ 5,  5,  2,  6,  4,  2,  7]])

assert_raises(IndexError,swap_cols,A,-1,1)
assert_raises(IndexError,swap_cols,A,7,1)
```

<!-- grade_id: cell-50441bfd51f5a7f6, points: 5, grade: true, solution: false -->
```python
# Check that the function `swap_cols` work as expected for a different matrix
# Marks: 5
```

**(c) [30 Marks]**

In this question you are asked to use the functions find_max_column_value, swap_rows, swap_columns (from parts (a) and (b)) together with another function add_row to implement the Gaussian Elimination algorithm using elementary row operations. 

The necessary functions will be available to your program and a brief description of how they work follows below. 

The specific implementation of Gaussian Elimination that you are asked to write down works as follows:

- Assume that $A=(a_{mn})$ is an $r \times k$ matrix given by a 2-dimensional NumPy array with entries of type `float64` 

- Then the Gauss Elimination algorithm can be described by the following steps:


1. Iterate over each column in turn. 
2. For each column, $j$, find the maximum absolute value below the diagonal and if larger than the diagonal value then swap the current row for the row corresponding to the maximum.
3. If the maximum value is 0 continue with next column.
4. Otherwise use elementary row operations to eliminate values in the current column from from all rows below the diagonal.
5. Once we have reached the end of the columns, return the current matrix. 


Additionally requirements:

- the function should raise a TypeError if the "dtype" is not float64
- To make the final output nicer you you should set all matrix element that satisfy  $|a_{mn}|<\min\{r,k\}\epsilon$ to $0$, 
where $\epsilon=1.19\cdot 10^{-7}$ (The number $\epsilon$ is sometimes called "machine-epsilon" and can be obtained in NumPy for single precision floating point numbers as `np.finfo('float32').eps`)

**Note:** For testing your function you can use some small matrices and work out the steps of the Gaussian elimination algorithm by hand. 

You can assume that the following functions are available (so you shouldn't write them again):
Here A is a matrix given by an NumPy array, i and j are integers and a is a float.  The functions all return an IndexError if i and j are not corresponding to rows or columns of A.

- `find_max_column_value(A,i,j)`: returns the  entry of A in column $j$ with maximum absolute value amongst rows $i,i+1,..$ together with its row index (counted in the entire matrix).
- `swap_rows(A,i,j)`: returns a copy of the matrix A where the  rows $i$ and $j$ are interchanged.
- `swap_columns(A,i,j)`: returns a copy of the matrix A where the  columns $i$ and $j$ are interchanged.
- `add_row(A,i,j,a)`:  returns a copy of A where row $$i$$ is replaced by row $i$ plus $a$ times row $j$
Since the function add_row is not part of the coursework we include the function definition below.


<!-- grade_id: cell-1f0b8c95d7800f5f, points: 0, grade: false, solution: true -->
```python
def add_row(A,i,j,a):
    if i < 0 or i >= A.shape[0]:
        raise IndexError(f"Index i={i} out of bounds!")
    if j < 0 or j >= A.shape[0]:
        raise IndexError(f"Index i={j} out of bounds!")
    B = A.copy()
    B[i] = A[i] + a*A[j]
    return B

def Gauss_elimination(A):
    # YOUR CODE HERE
    n = min(A.shape[0],A.shape[1])
    for i in range(A.shape[0]):
        for j in range(A.shape[1]):
            if type(A[i][j]) != np.float64:
                raise TypeError()
    for j in range(A.shape[1]):
        if find_max_column_value(A,min(A.shape[0],j),j)[0] > abs(A[j][j]):
            A = swap_rows(A,j,find_max_column_value(A,j,j)[1])
        if A[j][j] != 0:
            for i in range(j+1,A.shape[0]):
                A = add_row(A,i,j,(-1*(A[i][j]/A[j][j])))
                
    for x in range(A.shape[0]):
        for j in range(A.shape[1]):
            if abs(A[i][j]) < min(A.shape[0],A.shape[1])*np.finfo('float32').eps:
                A[i][j] = 0
    
    return A        
        
    raise NotImplementedError()
```

<!-- grade_id: cell-1f0b8c95d7800f5f, points: 0, grade: false, solution: true -->
```python
def add_row(A,i,j,a):
    if i < 0 or i >= A.shape[0]:
        raise IndexError(f"Index i={i} out of bounds!")
    if j < 0 or j >= A.shape[0]:
        raise IndexError(f"Index i={j} out of bounds!")
    B = A.copy()
    B[i] = A[i] + a*A[j]
    return B

def Gauss_elimination(A):
    # YOUR CODE HERE
    for i in range(A.shape[0]):
        for j in range(A.shape[1]):
            if type(A[i][j]) != np.float64:
                raise TypeError()
    for j in range(min(A.shape[0],A.shape[1])):
        if find_max_column_value(A,j,j)[0] > abs(A[j][j]):
            A = swap_rows(A,j,find_max_column_value(A,j,j)[1])
        if A[j][j] != 0:
            for i in range(j+1,A.shape[0]):
                A = add_row(A,i,j,(-1*(A[i][j]/A[j][j])))
                
    for x in range(A.shape[0]):
        for j in range(A.shape[1]):
            if abs(A[i][j]) < min(A.shape[0],A.shape[1])*np.finfo('float32').eps:
                A[i][j] = 0
    
    return A        
        
    raise NotImplementedError()
```

<!-- grade_id: cell-d1e587104abc3ea4, points: 5, grade: true, solution: false -->
```python
# Check that it works for a small example
# Marks: 5
from nose.tools import assert_almost_equal
import numpy as np
# Check that Echeblon form is correct
A = np.array([[1,2],[3,4]],dtype='float64')
A1 = np.array([[3. , 4. ], [0. , 2.0/3.0]])
eps = 2**2*np.finfo(A.dtype).eps
# Check that the Gauss elimination is calculated correctly to decent precision.
assert_equal( (abs(Gauss_elimination(A)-A1) < eps).all(),True)

```

<!-- grade_id: cell-af5e5e8bc57d1ed0, points: 5, grade: true, solution: false -->
```python
# Check that it works if one column becomes 0
# Marks: 5
from nose.tools import assert_equal
import numpy as np
eps = 2**2*np.finfo(A.dtype).eps
# Check that Echeblon form is correct
A = np.array([[1,2],[2,4]],dtype='float64')
A1 = np.array([[2., 4.], [0., 0.]])
assert_equal( (abs(Gauss_elimination(A)-A1) < eps).all(),True)


```

<!-- grade_id: cell-9d381a3116f0167e, points: 5, grade: true, solution: false -->
```python
# Check that the function return a type error if called with an array of different type
# Marks: 5
from nose.tools import assert_raises
import numpy as np
eps = 2**2*np.finfo(A.dtype).eps
# Check that Echeblon form is correct
A = np.array([[1,2],[2,4]],dtype='int')
assert_raises(TypeError,Gauss_elimination,A)

```

<!-- grade_id: cell-37ec76cd9d9f0669, points: 5, grade: true, solution: false -->
```python
# check that it works for a larger matrix
# Marks: 5
from nose.tools import assert_almost_equal
import numpy as np
# A larger test matrix A 
A = np.array([[ 9.,  6.,  7.,  8.,  1.,  7.,  2.],
       [ 8.,  2.,  6.,  5.,  1.,  5.,  3.],
       [ 7.,  3.,  1.,  4.,  5., 10.,  1.],
       [10.,  5.,  7.,  5.,  4.,  6.,  2.],
       [ 5.,  5.,  2.,  6.,  4.,  2.,  7.]])
# It's Echelon reduced form
A1 = np.array([[10.0, 5.0, 7.0, 5.0, 4.0, 6.0, 2.0], 
               [0.0, 2.5, -1.5, 3.5, 2.0, -1.0, 6.0], 
               [0.0, 0.0, -4.2, 1.2, 2.6, 5.6, 0.8], 
               [0.0, 0.0, 0.0, 3.5714285714285716, -1.0952380952380956, -1.6666666666666679, 6.047619047619048], 
               [0.0, 0.0, 0.0, 0.0, -2.24, 5.199999999999999, -6.239999999999999]])
eps = 2**7*np.finfo(A.dtype).eps
assert_equal( (abs(Gauss_elimination(A)-A1) < eps).all(),True)
```

<!-- grade_id: cell-1ba2e18248f7bd8d, points: 5, grade: true, solution: false -->
```python
# Check that it works for another matrix
# Marks: 5
from nose.tools import assert_almost_equal
import numpy as np

```

<!-- grade_id: cell-89c6d18565f2e5db, points: 5, grade: true, solution: false -->
```python
# check that it works for a column vector
# Marks: 5
from nose.tools import assert_almost_equal
import numpy as np

```